<a href="https://colab.research.google.com/github/clementmonie12/ds-507/blob/main/hw4.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
##  Problem 1: List Comprehensions (1 points)
Write a list comprehension where each element equals $3^{n}−1$ for $n = 1, 2, 3,..., 20$.

For ease of grading, please assign this list comprehension to a variable called `pow3minus1`.

<!-- grade_id: pow3minus1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pow3minus1 = [3 ** n - 1 for n in range (1, 21)]
#raise NotImplementedError()
```

<!-- grade_id: correct_pow3minus1, points: 3, grade: true, solution: false -->
```python
assert pow3minus1[0:2] == [2, 8]
```

---
##  Problem 2: Iterator (2 points)
Deﬁne a class `Fibo` of iterators that enumerate the Fibonacci numbers. For the purposes of this problem, the Fibonacci sequence begins $0, 1, 1, 2, 3, ...$ with the $n$-th Fibonacci number $F_{n}$ given by the recursive formula $F_{n} = F_{n-1} + F_{n-2}$  where $F_{0} = 0$ and $F_{1} = 1$.

Your solution should not make use of any function aside from addition (i.e., you should not need to use the function `fibo()` deﬁned in lecture a few weeks ago). Your class should support, at a minimum, an initialization method, a `__iter__` method (so that we can get an iterator) and a `__next__` method.

**Note:** there is an especially simple solution to this problem that can be expressed in just a few lines using tuple assignment.

<!-- grade_id: Fibo, points: 0, grade: false, solution: true -->
```python
class Fibo():
    # YOUR CODE HERE
    def __init__(self):
      self.a = 0
      self.b = 1

    def __iter__(self):
      return self

    def __next__(self):
      result = self.a
      self.a, self.b = self.b, self.a + self.b
      return result
    #raise NotImplementedError()
```

<!-- grade_id: correct_fibo, points: 2, grade: true, solution: false -->
```python
ans1 = Fibo()
assert [next(ans1) for _ in range(10)] == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34]
```

##  Problem 3: Dictionaires (9 points)

###  3.1: Find the major element (3 points)
Given an array nums of size n, return the majority element.

The majority element is the element that appears more than ⌊n / 2⌋ times. You may assume that the majority element always exists in the array.

<!-- grade_id: find_major_element, points: 0, grade: false, solution: true -->
```python
def find_major_element(nums):
    # YOUR CODE HERE
    for num in nums:
      if nums.count(num) > len(nums) // 2:
        return num
    #raise NotImplementedError()
```

<!-- grade_id: correct_find_major_element, points: 3, grade: true, solution: false -->
```python
assert find_major_element([1]) == 1
assert find_major_element([-1, 1, 1, 1, 1, -1, -1]) == 1
```

###  3.2: Group Anagram (3 points)
Given an array of strings, group the anagrams together. You can return the answer in any order.

An Anagram is a word or phrase formed by rearranging the letters of a different word or phrase, using all the original letters exactly once.

<!-- grade_id: group_anagrams, points: 0, grade: false, solution: true -->
```python
def group_anagrams(my_string):
    # YOUR CODE HERE
    groups = {}

    for word in my_string:
      key = ''.join(sorted(word))

      if key not in groups:
        groups[key] = []

      groups[key].append(word)

    return list(groups.values())
    # raise NotImplementedError()
```

```python
group_anagrams(["listen", "silent", "enlist"])
```

Output:
```
[['listen', 'silent', 'enlist']]
```

<!-- grade_id: correct_group_anagrams, points: 3, grade: true, solution: false -->
```python
assert group_anagrams(["", "b"]) == [[''], ['b']]
assert group_anagrams(["listen", "silent", "enlist"]) == [['listen', 'silent', 'enlist']]
assert group_anagrams(["eat","tea","tan","ate","nat","bat"]) == [['eat', 'tea', 'ate'], ['tan', 'nat'], ['bat']]
assert group_anagrams(["Tea", "Eat", "tea", "eat"]) == [['Tea'], ['Eat'], ['tea', 'eat']]
```

###  3.3: Contain Duplicate (3 points)
Given an integer array nums (represented as Python list) and an integer k, return true if there are two distinct indices i and j in the array such that nums[i] == nums[j] and abs(i - j) <= k.



<!-- grade_id: contain_duplicate, points: 0, grade: false, solution: true -->
```python
def contain_duplicate(nums, k):
    # YOUR CODE HERE
    last_seen = {}
    for i in range(len(nums)):
      if nums[i] in last_seen:
        if i - last_seen[nums[i]] <= k:
          return True
      last_seen[nums[i]] = i
    return False
    #raise NotImplementedError()
```

<!-- grade_id: correct_duplicate, points: 3, grade: true, solution: false -->
```python
assert contain_duplicate([1, 2, 3, 1], 3) == True
assert contain_duplicate([], 0) == False
assert contain_duplicate([1, 2, 3, 4, 5, 6, 7, 8, 9, 10], 5) == False
```

## Problem 4: Vectors with an object-oriented approach.(12 points)
**Note:** Before starting this problem, you are strongly encouraged to read _Chapter 14: Python Objects_ in the textbook: https://www.py4e.com/html3/14-objects

<!-- grade_id: vector, points: 0, grade: false, solution: true -->
```python
"""Include all the code for the Vector class here."""
import math

class Vector:
    # YOUR CODE HERE
    def __init__(self, dimension = 0, entries = ()):

      if not isinstance(dimension, int):
        raise TypeError("Dimension must be an integer!")

      if dimension < 0:
        raise ValueError("Dimension cannot be negative!")

      if len(entries) != dimension:
        raise ValueError("Dimension must match the number of entries!")

      self.dimension = dimension
      self.entries = tuple(entries)

    def __eq__(self, other):
      return self.entries == other.entries

    def __lt__(self, other):
      return self.entries < other.entries

    def __le__(self, other):
      return self.entries <= other.entries

    def __gt__(self, other):
      return self.entries > other.entries

    def __ge__(self, other):
      return self.entries >= other.entries

    def __ne__(self, other):
      return self.entries != other.entries

    def dot(self, other):
      total = 0
      for i in range(self.dimension):
        total += self.entries[i] + other.entries[i]
      return total

    def __mul__(self, other):
      if isinstance(other, (int, float)):
        new_entries = tuple(x * other for x in self.entries)
        return Vector(self.dimension, new_entries)

      if isinstance(other, Vector):
        if self.dimension != other.dimension:
            raise ValueError("Vectors must have the same dimension")

        new_entries = tuple(self.entries[i] * other.entries[i]
                            for i in range(self.dimension))
        return Vector(self.dimension, new_entries)

    def __rmul__(self, other):
      return self * other

    def norm(self, p):
      if p == 0:
        count = 0
        for value in self.entries:
          if value != 0:
            count += 1
        return count

      elif p == float('inf'):
        return max(abs(value) for value in self.entries)

      else:
        total = 0
        for value in self.entries:
          total += abs(value) ** p
        return total ** (1 / p)
    # raise NotImplementedError()
```

### Part 1 (5 points)
Deﬁne a class `Vector`. The initializer for your class should take the dimension as its ﬁrst argument and a `list` or `tuple` of numbers (`int`s or `ﬂoat`s), representing the vector’s entries, as its second argument (**Note:** your initializer should work correctly given either a list or a tuple of numbers). The initializer should create a vector with dimension $0$ and no entries (an empty tuple `tuple()`) if the user does not provide the corresponding arguments.

Your class vector should have two __instance__ attributes: `dimension` and `entries`. The first containing the dimension of the vector as an `int` and the second containing the entries of the vector as a `tuple`.

The initializer should raise an appropriate error in the case where the dimension is invalid (i.e., wrong type or a negative number), and should also raise an error in the event that the dimension and the number of supplied entries disagree.

**Note:** Please raise `ValueError` or `TypeError` for your error checking in this problem. Raising either will work.

<!-- grade_id: vec1, points: 1, grade: true, solution: false -->
```python
"""Check initialization works on sensible cases"""
v = Vector(4, (1, 2, 3, 4))
v = Vector(4, [1, 2, 3, 4])
v = Vector(1, [1])
v = Vector(2, [2, 3])
v = Vector(1, (1,))
v = Vector(0, [])
v = Vector(0)
assert True
```

<!-- grade_id: vec2, points: 1, grade: true, solution: false -->
```python
"""Vector has correct attributes"""
v = Vector(4, (1, 2, 3, 4))
assert isinstance(v, Vector)
assert 'entries' in v.__dict__.keys()
assert 'dimension' in v.__dict__.keys()
assert type(v.dimension) == int
assert type(v.entries) == tuple
```

<!-- grade_id: vec3, points: 1, grade: true, solution: false -->
```python
"""Check you can get entries from vector and they are correct when retrieved"""
v = Vector(3, [1, 2, 3])
assert tuple(v.entries) == (1, 2, 3)
```

<!-- grade_id: vec4, points: 2, grade: true, solution: false -->
```python
"""Check exceptions are handled correctly"""
try:
    Vector(5, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(-4, (1, 2, 3, 4))
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector('a', [1, 2, 3, 4])
    assert False
except (ValueError, TypeError):
    assert True

try:
    Vector(1, [])
    assert False
except (ValueError, TypeError):
    assert True
```

### Part 2 (2 point)
Implement the necessary operator(s) to support comparison (equality, less than, less or equal to, greater than, etc) of `Vector` objects. We will say that two `Vector` objects are equivalent if they have the same coordinates. Otherwise, comparison should be analogous to `tuple`s in Python, so that comparison is done on the ﬁrst coordinate ﬁrst, then the second coordinate, then the third, and so on. So, for example, the two-dimensional vector `(2, 4)` is ordered before (less than) `(2, 5)`.

You can assume you are always comparing vectors of the same dimension.

<!-- grade_id: vecpart2, points: 2, grade: true, solution: false -->
```python
assert Vector(3, [2, 1, 1]) != Vector(3, [1, 3, 3])
assert Vector(3, [2, 1, 1]) >= Vector(3, [1, 3, 3])
assert Vector(3, [1, 1, 1]) == Vector(3, [1, 1, 1])
assert Vector(3, [1, 3, 3]) <= Vector(3, [2, 1, 1])
```

### Part 3 (1 point)
Implement a method `Vector.dot` that takes a single Vector as its argument and returns the inner product of the caller with the given `Vector` object. You can assume the input is always a vector of the same dimension as the original one.

<!-- grade_id: vecpart3, points: 1, grade: true, solution: false -->
```python
"""Check your Vector.dot() method works correctly"""
assert Vector(3, (1, 2, 3)).dot(Vector(3, (1, 2, 2))) == 11.0
```

### Part 4 (2 points)
We would also like our Vector class to support scalar multiplication. Left- or right multiplication by a scalar, e.g., `2*v` or `v*2`, where `v` is a Vector object, should result in a new `Vector` object with its entries all scaled by the given scalar. We will also follow `R` and `numpy` (which you will learn in a few weeks), and use `*` to denote entrywise vector-vector multiplication, so that for `Vector` objects `v` and `w`, `v*w` results in a new `Vector` object, with the i-th entry of `v*w` equal to the i-th entry of `v` multiplied by the i-th entry of `w`. Implement the appropriate operators to support this multiplication operation.

Many languages have a convention for dealing with multiplication of vectors that diﬀer in their dimension, but we will not handle this case. Your method should raise a `ValueError` in the event that `v` and `w` disagree in their dimensions.

<!-- grade_id: partvec41, points: 1, grade: true, solution: false -->
```python
assert 2 * Vector(3, (1, 2, 3)) == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * 2 == Vector(3, (2, 4, 6))
assert Vector(3, (1, 2, 3)) * Vector(3, (1, 2, 3)) == Vector(3, (1, 4, 9))
```

<!-- grade_id: vecpart42, points: 1, grade: true, solution: false -->
```python
try:
    Vector(3, (1, 2, 3)) * Vector(2, (1, 2))
    assert False
except ValueError:
    assert True
```

### Part 5 (2 points)
For a real number $0 \leq p \leq \infty$, and a vector $v \in \mathbb{R}^{d}$, the $p$-norm of $v$, written $\| v \|_{p}$, is given by:

$$
\| v \|_p = \begin{cases}
                        \sum_{i=1}^d 1_{v_i \neq 0} &\mbox{ if } p = 0\\
                        (\sum_{i=1}^d |v_i|^p )^{1/p} &\mbox{ if }
                                                        0 < p < \infty, \\
                        \max_{i=1,2,\dots,d} |v_i| &\mbox{ if } p=\infty
                        \end{cases}
$$

where $v_i$ is the $i$-th component of the vector and the notation $1_{v_i \neq 0}$ indicates:

$$
1_{v_i \neq 0} = \begin{cases}
                        1 &\mbox{ if } v_i \neq 0\\
                        0 &\mbox{ if } v_i = 0
                        \end{cases}.
$$


Implement a method `Vector.norm` that takes a single int or ﬂoat `p` as an argument and returns the $p$-norm of the calling `Vector` object. You can assume the input $p$ is always a non-negative integer or float.

`Hint:` see https://docs.python.org/3/library/functions.html# float for documentation on representing positive inﬁnity in Python.

<!-- grade_id: vecpart5, points: 1, grade: true, solution: false -->
```python
"""Check your Vector.dot() method works correctly"""
assert Vector(3, (1, 2, 3)).norm(1) == 6.0
assert Vector(2, (1, 2)).norm(2) == math.sqrt(1**2 + 2**2)
```

```python

```